# Assignment 1

## 1. (1.0 point) Function return value

This exercise is to show you how function return values work. Run the code below to see what happens.

```python
def f1(x):
    print(x + 1)

def f2(x):
    return(x + 1)

f1(4)
f2(4)
```

Output:
```
5
5
```

Each function should produce the same output. Now run the following code.

```python
f1(4) + 1
f2(4) + 1
```

Output:
```
5
TypeError         Traceback (most recent call last)
<ipython-input-2-02362f8b3abb> in <module>()
----> 1 f1(4) + 1
      2 f2(4) + 1

TypeError: unsupported operand type(s) for +: 'NoneType' and 'int'
```

Explain what is happening and why you get this result.

f2 returns a value - in this case 5, which is an integer. Hence we can add another integer to it.

f1, on the other hand, returns a value None - because we didn't provide a return statement explicitly. 
The print statement does not return a value to the function - instead, it just shows a value on the screen. It
does not save the value for the program.

## 2. (1.0 point) Function operation

Consider the following piece of code:

```python
def f(x):
    if x == 0 or x == 1:
        return x
    return f(x - 1) + f(x - 2)
```

Describe, in words, what this code does, and how it does it.

This is a recursive function that reduces a value by 1.

First it takes a value and checks if it is equal to 0 or 1. If it is equal to 0, it will return 0. If it is equal to 1, it will return 1.

If the value is not 0 nor 1, it will calculate f(x-1) + f(x-2), which is recursion. In case of 2, the output will be 1. In case of 3, the output will be 2. In case of 4, the output is 4, etc.

For what inputs will this function not behave as expected? What will happen?

For negatie values, this function will not converge as the values will continue getting more negative. 
This function also will not work for some data types, such as Strings - we would receive an error.

## 3. (1.0 point) Defining a function

The following program plays the game of "Rock, Paper, and Scissors". Familiarize yourself with how it works. Try running the program.

```python
from random import randint
valid_input = ["R", "P", "S"]
winners = [["R", "S"], ["S", "P"], ["P", "R"]]

print("You are playing a Rock, Paper, Scissors simulation.")
user_choice = input("Select R for rock, P for Paper, S for scissors, Q to quit\n> ")
user_choice = user_choice.upper()

while user_choice != 'Q':
    if user_choice not in valid_input:
        print("Error: invalid input -- " + user_choice)
        break
    
    computer_choice = valid_input[randint(0, len(valid_input)-1)]
    print("Computer chooses " + computer_choice)

    if user_choice == computer_choice:
        print("Tie Game")
    elif [user_choice, computer_choice] in winners:
        print("You win!")
    else:
        print("Computer wins!")
    
    user_choice = input("Select R for rock, P for Paper, S for scissors, Q to quit\n> ")
    user_choice = user_choice.upper()
```

Output:
```
You are playing a Rock, Paper, Scissors simulation.
Select R for rock, P for Paper, S for scissors, Q to quit
> R
Computer chooses P
Computer wins!
Select R for rock, P for Paper, S for scissors, Q to quit
> P
Computer chooses R
You win!
Select R for rock, P for Paper, S for scissors, Q to quit
> S
Computer chooses P
You win!
Select R for rock, P for Paper, S for scissors, Q to quit
> Q
```

Recall how we define a function using `def`, and how we pass in parameters. Transform the code into a function called `rps` that takes two parameters: `user_choice` and `computer_choice`, with valid inputs `'R'`, `'P'`, and `'S'`. Your function should return `'D'` for a draw, `'W'` if `user_choice` wins, and `'L'` if `user_choice` loses.

```python
### BEGIN SOLUTION
def rps(user_choice, computer_choice):
    ''' Function that takes two parameters: user_choice and computer_choice, with valid inputs R, P, S.
    It returns D for a draw, W if user wins, and L if user loses.'''
    
    winners = [["R", "S"], ["S", "P"], ["P", "R"]]
    
    if user_choice == computer_choice:
        return 'D'
    elif [user_choice, computer_choice] in winners:
        return 'W'
    else:
        return 'L'
    
### END SOLUTION
```

<!-- grade_id: Q3C, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(rps('R', 'R'), 'D')
assert_equal(rps('R', 'P'), 'L')
assert_equal(rps('R', 'S'), 'W')
assert_equal(rps('P', 'R'), 'W')
assert_equal(rps('P', 'P'), 'D')
assert_equal(rps('P', 'S'), 'L')
assert_equal(rps('S', 'R'), 'L')
assert_equal(rps('S', 'P'), 'W')
assert_equal(rps('S', 'S'), 'D')
```

## 4. (1.0 point) List function

Given a list of non-empty tuples, return a list sorted in increasing order by the last element in each tuple. 

For example: `[(1, 7), (1, 3), (3, 4, 5), (2, 2)]` yields `[(2, 2), (1, 3), (3, 4, 5), (1, 7)]`

<!-- grade_id: Q4A, points: 0, grade: false, solution: true -->
```python
def sort_tuple_list(tuple_list):
    '''Function to sort a list of non-empty tuples in increasing order by the last element in each tuple'''
    ### BEGIN SOLUTION
    ### END SOLUTION
```

<!-- grade_id: Q4B, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(sort_tuple_list([(1, 3), (3, 2), (2, 1)]), [(2, 1), (3, 2), (1, 3)])
assert_equal(sort_tuple_list([(2, 3), (1, 2), (3, 1)]), [(3, 1), (1, 2), (2, 3)])
assert_equal(sort_tuple_list([(1, 7), (1, 3), (3, 4, 5), (2, 2)]), [(2, 2), (1, 3), (3, 4, 5), (1, 7)])
```

# 5. (1.0 point) List sort

Given two lists sorted in increasing order, create and return a merged list of all the elements in sorted order. You may modify the passed in lists. Ideally, the solution should work in "linear" time, making a single pass of both lists.

<!-- grade_id: Q5A, points: 0, grade: false, solution: true -->
```python
def list_merge(list1, list2):
    '''Given two lists sorted in increasing order, create and 
       return a merged list of all the elements in sorted order
    '''
    ### BEGIN SOLUTION
    ### END SOLUTION
```

<!-- grade_id: Q5B, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(list_merge(['aa', 'xx', 'zz'], ['bb', 'cc']), ['aa', 'bb', 'cc', 'xx', 'zz'])
assert_equal(list_merge(['aa', 'xx'], ['bb', 'cc', 'zz']), ['aa', 'bb', 'cc', 'xx', 'zz'])
assert_equal(list_merge(['aa', 'aa'], ['aa', 'bb', 'bb']), ['aa', 'aa', 'aa', 'bb', 'bb'])
```

# 6. (1.0 point) Built-in Operator

Imagine that Python doesn’t have the `!=` operator built in. Write a function that takes two parameters and gives the same result as the `!=` operator. Obviously, you cannot use `!=` within your function! Test if your code works by thinking of examples and making sure the output is the same for your new method as `!=` gives you.

<!-- grade_id: Q6A, points: 0, grade: false, solution: true -->
```python
def not_equal(val1, val2):
    ''' User supplied version of the != operator, implemeted without using !='''
    ### BEGIN SOLUTION
    ### END SOLUTION
```

<!-- grade_id: Q6B, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_true, assert_false
assert_true(not_equal(2,3))
assert_false(not_equal(3,3))
assert_true(not_equal('3',3))
assert_true(not_equal((1,2),[1,2]))
assert_false(not_equal((3,4),(3,4)))
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-3-d9667c7ec52c> in <module>()
      1 from nose.tools import assert_true, assert_false
----> 2 assert_true(not_equal(2,3))
      3 assert_false(not_equal(3,3))
      4 assert_true(not_equal('3',3))
      5 assert_true(not_equal((1,2),[1,2]))

NameError: name 'not_equal' is not defined
```

## 7. (1.5 point) Quadratic Formula

A quadratic equation is a second order polynomial with three coefficients a, b, and c. It is given by:

$$ax^2 + bx + c = 0$$

The solutions to the equation are given by the formulas:

$$x_1 = \frac{-b + \sqrt{b^2 - 4ac}}{2a}$$

and

$$x_2 = \frac{-b - \sqrt{b^2 - 4ac}}{2a}$$



Write a function that computes the roots of a quadratic equation and returns them as a tuple. Your function should take three parameters, a, b, and c. Note that if

$$b^2 - 4ac < 0$$ 

then the roots will be complex. Hint: remember that Python has a built-in complex data type.

<!-- grade_id: Q7A, points: 0, grade: false, solution: true -->
```python
def roots(a, b, c):
    ''' Computes the roots of the quadratic equation with coefficients a, b, c, and returns
        them as the tuple (x1, x2)
    '''
    ### BEGIN SOLUTION
    ### END SOLUTION
```

<!-- grade_id: Q7B, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(roots(1,2,3), ((-1+1.4142135623730951j), (-1-1.4142135623730951j)))
assert_equal(roots(0,0,4), (0j, 0j))
assert_equal(roots(-4,3,2), (-0.42539052967910607, 1.175390529679106))
```

## 8. (2.5 points) The Game of Nims

In this game, two players sit in front of a pile of 100 stones. They take turns, each removing between 1 and 5 stones (assuming there are at least 5 stones left in the pile). The person who removes the last stone(s) wins.

In this problem, you’ll write a function to play this game. It may seem tricky, so break it down into parts. Like many programs, we have to use nested loops (one loop inside another). In the outermost loop, we want to keep playing until we are out of stones. Inside that, we want to keep alternating players. You have the option of either writing two blocks of code, or keeping a variable that tracks the current player.

We are going to use the `input` function to obtain the user's move. We might want to have an innermost loop that checks if the user’s input is valid. Is it a number? Is it a valid number (e.g. between 1 and 5)? Are there enough stones in the pile to take off this many? If any of these answers are no, we should tell the user and re-ask them the question.

If you choose to write two blocks of code, the basic outline of the program should be something like this:

    while [pile is not empty]:
       while [player 1’s answer is not valid]:
          [ask player 1]
       [execute player 1’s move]
       [same as above for player 2]
       
Be careful with the validity checks. Specifically, we want to keep asking player 1 for their choice as long as their answer is not valid, BUT we want to make sure we ask them at least ONCE. So, for example, we will want to keep a variable that tracks whether their answer is valid, and set it to `False` initially.

There is no need to limit yourself to only one function! Use as many functions as you want in order to make the code more readable and modular.

When you're finished, test each other's programs by playing them!

<!-- grade_id: Q8A, points: 4, grade: true, solution: true -->
```python
def play_nims(pile, max_stones):
    '''
    An interactive two-person game; also known as Stones.
    @param pile: the number of stones in the pile to start
    @param max_stones: the maximum number of stones you can take in one turn
    '''

## Basic structure of program (feel free to alter as you please):

#    while [pile is not empty]:
#        while [player 1's answer is not valid]:
#            [ask player 1]
#            [execute player 1's move]
#       
#        while [player 2's answer is not valid]:
#            [ask player 2]
#            [execute player 2's move]
#
#    print "Game over"
### BEGIN SOLUTION
### END SOLUTION
```

```python

```